# Lab 1b
CS1302 Introduction to Computer Programming
___

The goal of this lab is to let you get a basic understanding of how to program. Don't worry if you haven't programmed before and don't get enough time to review Lectures.

This Lab sheet involves exercises in 2 parts: Guided Exercises (Q1-Q7) --> Skill-Drilling Exercise (Q8-Q10) 

Hint: the output samples can be found [here](https://www.cs.cityu.edu.hk/~helena/cs1302202021B/bin/Lab01b_samples.html) <span style="color:#ff0000"><== You may copy the content directly to avoid some typing mistakes</span>

## Guided Exercises (Q1-Q7)
### Q1. The simpliest Hello World program (Version 1)
As mentioned in Lab 1, the first program for a beginner to write is often the ["Hello, World!"](https://en.wikipedia.org/wiki/%22Hello,_World!%22_program) program, which says Hello to the world. Please write your program below which prints **Hello, World!**. If you forget how to do so, please review Section 2.1 of setup.ipynb in Lab 1.  Remember to always delete *raise NotImplementedError()* before typing any code.

<!-- grade_id: Lab1bQ1HelloWorld, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
print('Hello, World!')
### END SOLUTION

```

The following is **a visible test** you can run to check your answer: The test returns an assertion error only if your program does not print the correct message.

<!-- grade_id: Lab1bQ1HelloWorldTesting, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your Q1 program.
import sys, io
old_stdout, sys.stdout = sys.stdout, io.StringIO()
exec(In[-2])
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == 'Hello, World!\n'
```

### Q2. Fix a given Hello World program (Version 2)

A program is given as comment in the code block below, with punctuations and spaces missing. Please fix it and make sure it can print the following contents:

<center><img src='1.png', width='25%''></center>
    
Note: 
   - In Python, you must use single quote, double quote or triple quote to enclose the string you want to print. Hence, *print('Hello, World!')*, *print("Hello, World!")*, or *print('''Hello, World!''')* is correct, but *print(Hello, World!)* is wrong.
   - For this question, the width of the message is 13, meaning each line contains 13 symbols/characters/punctuations in total.

<!-- grade_id: Lab1bQ2Dashes, points: 0, grade: false, solution: true -->
```python
# print-------------
# printHello, World!
# print-------------

### BEGIN SOLUTION
print('-------------')
print('Hello, World!')
print('-------------')
### END SOLUTION
```

<!-- grade_id: Lab1bQ2DashesTesting, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your Q2 program.
import sys, io
old_stdout, sys.stdout = sys.stdout, io.StringIO()
exec(In[-2])
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == '''-------------\nHello, World!\n-------------\n'''
```

### Q3. Create another Hello World program (Version 3)

[Escape sequences](https://en.wikipedia.org/wiki/Escape_sequence) allow you to include special characters in strings. To do this, simply add a backslash (\\) before the character you want to escape. Click [here](https://www.w3schools.com/python/gloss_python_escape_characters.asp) for more information. Please try the code below to help you understand how to print special symbols in Python.
<center><img src='0.png', width='50%''></center>

```python
# this code shows how to use escape sequence to print some special symbols
# this is how you print a single quote
print('\'')
# this is how you print a backslash
print('\\')

# for some symbols, you can print them directly
print('/')
print('?')
print('|')
```

After understanding how to print special symbols, you need to write a program to print the following pattern:
<center><img src='2.png', width='25%''></center>
    
Hint: 
   * Use one *print()* function to print each line. So you need 5 *print()* functions in total.
   * You can first print the middle line: "| Hello, World! |", then add the top and bottom line.
   * The total length of each line is 16.
   * For the vertical stroke ```|```, you can find it on your keyboard as shown [here](https://www.cs.cityu.edu.hk/~helena/misc/bin/vertical_stroke.png)
  


<!-- grade_id: Lab1bQ3Frame, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
print('/---------------\\')
print('|               |')
print('| Hello, World! |')
print('|               |')
print('\\---------------/')
### END SOLUTION
```

<!-- grade_id: Lab1bQ3FrameTesting, points: 1, grade: true, solution: false -->
```python
# Run this test cell right after running your Q3 program.
import sys, io
old_stdout, sys.stdout = sys.stdout, io.StringIO()
exec(In[-2])
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
assert printed == '''/---------------\\\n|               |\n| Hello, World! |\n|               |\n\\---------------/\n'''
```

### Q4. Welcome World or Hello World
The following program is to ask the user for his choice: 'w' or 'h', and show the corresponding message.

* If he inputs 'w', print 'Welcome, World!'
* If he inputs 'h', print 'Hello, World!' 
* For other cases, display nothing:

Your task: Study the code and add one line of statement to complete it.

Note also:  
1. [Indentation](https://www.w3schools.com/python/gloss_python_indentation.asp#:~:text=Indentation%20refers%20to%20the%20spaces,indicate%20a%20block%20of%20code.) is very important in Python. It is usually a Tab key or four spaces.  
2. DO NOT USE **else** for this question

<!-- grade_id: Lab1bQ4WH, points: 1, grade: true, solution: true -->
```python
choice=input("Input your choice ('w' or 'h'): ") # A variable is defined. Its name is choice
#An Input-Statement. It waits for the user's input (ended with <enter>), and stores in the variable choice. 

#This is an if-statement, if the condition is met, the code below it will be executed; otherwise, the code is skipped
if choice=='w':
    print('Welcome, World!') #note the indentation before print()
if choice=='h':    
    ### BEGIN SOLUTION
    print('Hello, World!')
### END SOLUTION


```

### Q5. Welcome World or Hello World
Based on the style of Q4, we need to get user's input '1', '2', or '3' (using the choice variable) and show the right version of "Hello, World!". 
* If the user inputs '1', print 'Hello, World!' in style of Q1
* If the user inputs '2', print 'Hello, World!' in style of Q2
* If the user inputs '3', print 'Hello, World!' in style of Q3
* If the user types a wrong choice, the program should show nothing 

Note: 
1. For each case, you may simply copy your code from previous questions.
1. DO NOT USE **else** for this question


<!-- grade_id: Lab1bQ5123, points: 1, grade: true, solution: true -->
```python
choice=input("Input your choice ('1', '2', or '3'): ") 
### BEGIN SOLUTION
if choice=='1':
    print('Hello, World!') #note the indentation before print()
if choice=='2':
    print('-------------')
    print('Hello, World!')
    print('-------------')
if choice=='3':
    print('/---------------\\')
    print('|               |')
    print('| Hello, World! |')
    print('|               |')
    print('\\---------------/')
### END SOLUTION
```

### Q6. Hello or Welcome - 1, 2, 3 
Based on Q4-Q5, create an extended program which lets the user choose both the message and the version

<center><img src='3.png', width='50%'></center>

Note:  DO NOT USE **else** for this question

<!-- grade_id: Lab1bQ6wh123, points: 2, grade: true, solution: true -->
```python
choiceWH=input("Input your choice for the message ('w' or 'h'): ") 
choice123=input("Input your choice for the style ('1', '2', or '3'): ") 

if choiceWH=='w':
    if choice123=='1':
        print("Welcome, World!")
    if choice123=='2':
        ### BEGIN SOLUTION
        print('---------------')
        print('Welcome, World!')
        print('---------------')
    if choice123=='3':
        print('/-----------------\\')
        print('|                 |')
        print('| Welcome, World! |')
        print('|                 |')
        print('\\-----------------/')
if choiceWH=='h':
    if choice123=='1':
        print('Hello, World!') #note the indentation before print()
    if choice123=='2':
        print('-------------')
        print('Hello, World!')
        print('-------------')
    if choice123=='3':
        print('/---------------\\')
        print('|               |')
        print('| Hello, World! |')
        print('|               |')
        print('\\---------------/')
### END SOLUTION
      
```

### Q7. Calculation of Study Hours

The following program calculates study hours based on the credit units of a course:


```python
x=input('Input the number of credit units of this course: ') # input 3 when a window is prompted
x=int(x) #this line is used to convert a character to an integer. Because the return data of input() is a character
print('In this semester, you need to study approximately',40*x,'to',50*x, 'hours for this course')
```

Your task: modify the program to show the calculation details like

<center><img src='4.png', width='90%''></center>
    
Hint: ```print('40 x', x,'=',40*x)``` gives ```40 x 3 = 120```

<!-- grade_id: Lab1bQ7StudyHrs, points: 2, grade: true, solution: true -->
```python
x=input('Input the number of credit units of this course: ') # input 3 when a window is prompted
x=int(x)
### BEGIN SOLUTION
print('40 x', x,'=',40*x)
print('50 x', x,'=',50*x)
print('In this semester, you need to study approximately',40*x,'to',50*x, 'hours for this course')
### END SOLUTION
```

## Skill-Drilling Exercise (Q8-Q10) 
### Q8. Greeting
Create a program like Q6, but this time ask the user to choose among 2 styles only (please refer to Q1 and Q2), while 3 messages will be available: 'w' for "Welcome!", 'h' for "Hello!", and 'c' for "Cheers!".

Example 1:
<center><img src='5.png', width='60%''></center>

Example 2:    
<center><img src='5-2.png', width='60%''></center>

<!-- grade_id: Lab1bQ812whc, points: 2, grade: true, solution: true -->
```python
choice12=input("Input your choice for the style ('1' or '2'): ") 
choiceWHC=input("Input your choice for the message ('w','h', or 'c'): ") 
### BEGIN SOLUTION
if choice12=='1':
    if choiceWHC=='w':
        print('Welcome!') #note the indentation before print()
    if choiceWHC=='h':
        print('Hello!')
    if choiceWHC=='c':
        print('Cheers!')

if choice12=='2':
    if choiceWHC=='w':
        print('--------')
        print('Welcome!')
        print('--------')
    if choiceWHC=='h':
        print('------')
        print('Hello!')
        print('------')
    if choiceWHC=='c':
        print('-------')
        print('Cheers!')
        print('-------')
### END SOLUTION
```

### Q9. Booking Day Time 
Imagine that in order to book a venue in the current week, we need to state the day and the duration (start hour and end hour).  Create a program as shown below, which prompts for these inputs and summarizes the booking in one line.
<center><img src='6.png', width='70%''></center>
    
Hint: we assume that the end hour is always larger than the start hour. Moreover, the duration is at least 2 hours. So we don't consider special cases such as start hour is 8 while the end hour is 6.


<!-- grade_id: Lab1bQ9booking, points: 2, grade: true, solution: true -->
```python
day=input("Input the day ('M'-Mon, 'T'-Tue, 'W'-Wed, 'R'-Thu, 'F'-Fri): ") 
h1=input("Input the start hour: ")
h2=input("Input the end hour: ")
### BEGIN SOLUTION
if day=='M':
    day_print='Monday'
if day=='T':
    day_print='Tuesday'
if day=='W':
    day_print='Wednesday'
if day=='R':
    day_print='Thursday'
if day=='F':
    day_print='Friday'

duration=int(h2)-int(h1)

print('')
print('Your booking is:',day_print, h1,'-',h2,'(Total:',duration,'hours)')
### END SOLUTION
```

### Q10. Print misteries
Create a program which displays the contents as shown below (no user input).

<center><img src='7.png', width='50%''></center>
    
Hint:   
1. To add an empty line, you can use ```print('')```, which means print nothing in this line.  
2. The correct output is [here](https://www.cs.cityu.edu.hk/~helena/cs1302202021B/bin/Lab01b_samples.html) <span style="color:#ff0000"> <== You may copy the content directly to avoid some typing mistakes.</span>

<!-- grade_id: Lab1bQ10Learning, points: 0, grade: false, solution: true -->
```python
print('Peter has learnt the following from lab01: ')
### BEGIN SOLUTION
print('')
print('''(1) It is wrong to write: print(Hello, World!);
    It is right to write: print(\'Hello, World!\');''')
print('')
print('''(2) To output one \'\\\', we need to type \'\\\\\'.
    Example: print(\'\\\\\')''')
### END SOLUTION
```

<!-- grade_id: Lab1bQ10LearningTesting, points: 2, grade: true, solution: false -->
```python
# Run this test cell right after running your Q10 program.
import sys, io
old_stdout, sys.stdout = sys.stdout, io.StringIO()
exec(In[-2])
printed = sys.stdout.getvalue()
sys.stdout = old_stdout
ans='''Peter has learnt the following from lab01: (1) It is wrong to write: print(Hello, World!);It is right to write: print(\'Hello, World!\');(2) To output one \'\\\', we need to type \'\\\\\'.    Example: print(\'\\\\\')'''
assert ''.join(printed.split()) == ''.join(ans.split())
```

-- END OF EXERCISES --